A stochastic explanation for observed local-to-global foraging states in Caenorhabditis elegans

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #2 (Public reviews):

Weaknesses:

This manuscript has two weaknesses that dampen the enthusiasm for the results. First, in all of the examples the authors cite where a Gillespie algorithm is used to sample from a distribution, be it the kinetics associated with chemical dynamics, or a Lotka-Volterra Competition Model, there are underlying processes that govern the evolution of the dynamics, and thus the sampling from distributions. In one of their references for instance, the stochasticity arises from the birth and death rates, thereby influencing the genetic drift in the model. In these examples, the process governing the dynamics (and thus generating the distributions from which one samples) are distinct from the behavior being studied. In this manuscript, the distribution being sampled from is the exponential decay function of the reorientation rate (lines 100-102). This appears to be tautological - a decay function fitted to the reorientation data is then sampled to generate the distributions of the reorientation data. That the model performs well, and matches the data is commendable, but it is unclear how that could not be the case if the underlying function generating the distribution was fit to the data.

To use the Lotka-Volterra model as an analogy, the changing reorientation rate (like a changing rate of prey growth) is tied to the decay in M (like a loss of predators). You could infer the loss of predators by measuring the changing rate of prey growth. In our case, we infer the loss of M by observing the changing reorientation rate. In the LotkaVolterra model, the prey growth rate is negatively associated with predator numbers, but in our model, the reorientation rate is positively associated with M, hence a loss in M leads to a decay in the reorientation rate.

You are correct that the decay parameters fit to the average should produce a distribution of in silico data that reproduce this average result (Figure 3a). However, this does not necessarily mean that these kinetic parameters should produce the same distributions of switch kinetics observed in Figure 3b. Indeed, a binary model (𝑴 ∈ {𝟎, 𝟏}), which produces an average distribution that matches the average experimental data (Figure 3a) produces a fundamentally different (bimodal) distribution of switch distributions in Figure 3b.

The second weakness is somewhat related to the first, in that absent an underlying mechanism or framework, one is left wondering what insight the model provides. Stochastic sampling a function generated by fitting the data to produce stochastic behavior is where one ends up in this framework, and the authors indeed point this out: "simple stochastic models should be sufficient to explain observably stochastic behaviors." (Line 233-234). But if that is the case, what do we learn about how the foraging is happening. The authors suggest that the decay parameter M can be considered a memory timescale; which offers some suggestion, but then go on to say that the "physical basis of M can come from multiple sources". Here is where one is left for want: The mechanisms suggested, including loss of sensory stimuli, alternations in motor integration, ionotropic glutamate signaling, dopamine, and neuropeptides are all suggested: this is basically all of the possible biological sources that can govern behavior, and one is left not knowing what insight the model provides. The array of biological processes listed are so variable in dynamics and meaning, that their explanation of what govern M is at best unsatisfying. Molecular dynamics models that generate distributions can point to certain properties of the model, such as the binding kinetics (on and off rates, etc.) as explanations for the mechanisms generating the distributions, and therefore point to how a change in the biology affects the stochasticity of the process. It is unclear how this model provides such a connection, especially taken in aggregate with the previous weakness.

Providing a roadmap of how to think about the processes generating M, the meaning of those processes in search, and potential frameworks that are more constrained and with more precise biological underpinning (beyond the array of possibilities described) would go a long way to assuaging the weaknesses.

The insight we (hopefully) are trying to convey is that individual observations of apparent state-switching behavior does not necessarily imply that a state change is actually happening if a large fraction of the population is not producing this behavior. This same observation can be recreated by invoking a stochastic process, which we already know is how reorientation occurrences behave in the first place. Apparent switches to global foraging are simply due to the reorientation rate decaying in time, not necessarily due to a sudden state change. We modeled a stochastic binary switch (when M0=1) which produced a bimodal distribution of switch kinetics (Figure 3b) which was different than the experimental distribution. The biological basis of M is not addressed here, but we clarified the language on lines 342 and 343 to reinforce that it likely represents the timescales of AIA and ADE activities. We reiterated what was described in López-Cruz et al to convey that molecularly, what is governing the timescales of these two neurons is not trivial, and likely multi-faceted.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

The presentation of the Gillespie algorithm, though much improved, is tough going and for many biologists will be a barrier to appreciation of what was done and what was achieved. I found the description of the algorithm generated by AI (ChatGTP) to be more accessible and the example given to be better related to the present application of the algorithm. This might provide a template for a more accessible description of the model.

We are glad the newer draft is clearer, and apologize it is still difficult to read. We made a few changes that hopefully clarify some points (see below).

It is unclear how instances of >1 transition were automatically distinguished from instances with 1 transition. A related point is how the transition-finding algorithm was kept from detecting too many transitions, as it seems that any quadruplet of points defines a slope change.

In López-Cruz et al, >1 transitions (and all transitions) were distinguished by eye after running the findchangepts function. We added a clarifying statement on lines 78 and 79 to illuminate this point. As noted on line 72, the function itself only fits two regressions, so by definition, it can only define one transition. This is why we decided to plot the distribution of slope and transition parameters in the first place; to see if there was a clear bimodal distribution (as observed for other observably binary states, like roaming and dwelling). This was not the case for the experimental data, but was observed in the in silico data if we forced the algorithm to be a two-state model (Figure 3b, M0 = 1).

Line 113-4: I was confused by the distinction between the probability of observing an event and the propensity for it to occur. Are the authors implying that some events occur but are not observed?

We apologize for this confusion, and added some phrasing in Lines 115-130 to address this. The propensity is analogous to the rate of a reaction. Given this rate, the probability of seeing Ω+1 reorientations in the infinitesimal time interval dt is the product of the propensity and the probability the current state is Ω reorientations.

Line 120: Shouldn't propensity at t = 0 be alpha + beta?

Yes, thank you for catching this. We fixed it.

Why was it necessary to posit two decay processes (equations 2 and 5?). Wouldn't one suffice?

Thank you, we have added some text to clarify this point (lines 129-132). The Gillespie algorithm models discrete temporal events, which are explicitly dependent on the current state of the system. Since the propensity itself is changing in time, it implies that it is coupled to another state variable that is changing in time, i.e. another propensity. Since an exponential decay is sufficient to model the decay in reorientations, this implies that the reorientation propensity is coupled to a first order decay propensity (equations 4-5).

Line 145: ...sudden changes in [reorientation rate] are not due to...

Thank you, we have corrected this (Line 157).

Fig. 2d: Legend implies (but fails to state) that each dot is a worm, raising the question of how single worms with multiple transitions were plotted in this graph as they would have more than one transition point.

Thank you, we updated the legend. Multiple transitions are not quantified with the tworegression approach. Prior observations, such as by López-Cruz, were simply done by eye.

Line 153: Does i denote either process 1 or 2?

Yes, i is the subscript for each propensity ai. We have added text on line 166 to clarify this.

Line 159: Confusing. If an "event" is a reorientation event and a "transition" is a discrete change in slope of Omega vs t, then "The probability that no events will occur for ALL transitions in this time interval" makes no sense.

Thank you, we have reworded this part (Lines 169-172) to be clearer.

Equation 17:Unclear what index i refers to

Thank you, we have changed this to index to j, and modified the text on line 228 to reflect this.

Line 227-9: Unclear how collisions are thought to have caused the shift in experimental distribution.

We have clarified the text on lines 246 and 250. Collisions are not being referred to here, but instead the crossing of pheromone trails. This is purely speculative.

Line 310-317. If M rises on food, then worms should reorient more on food than after long times off food, when M has decayed. But worms don't reorient much on food; they behave as though M is low. This seems like a contradiction, unless one supposes instead that M is low on food and after long times off food but spikes when food is removed.

Thank you, we have added clarifying language on lines 333-336 to address this point. Worm behavior is fundamentally different on food, as worms transition to a dwell/roam behavioral dynamic which is fundamentally different than foraging behavior while off food.

Changes from the earlier version to the revised version

Abstract

Abstract

Abrupt changes in behavior can often be associated with changes in underlying behavioral states. When placed off food, the foraging behavior of C. elegans can be described as a change between an initial local-search behavior characterized by a high rate of reorientations, followed by a global-search behavior characterized by sparse reorientations. This is commonly observed in individual worms, but when numerous worms are characterized, only about half appear to exhibit this behavior. We propose an alternative model that predicts both abrupt and continuous changes to reorientation that does not rely on behavioral states. This model is inspired by molecular dynamics modeling that defines the foraging reorientation rate as a decaying parameter. By stochastically sampling from the time interval probability distribution defined by this rate, both abrupt and gradual changes to reorientation rates can occur, matching experimentally observed results. Crucially, this model does not depend on behavioral states or information accumulation. Even though abrupt behavioral changes do occur, they are not necessarily indicative of abrupt changes in behavioral states, especially when abrupt changes are not universally observed in the population.

Introduction

The search for food in the absence of informative sensory cues is an essential animal behavior1. A foraging strategy that is observed in nearly all animals is the Area Restricted Search (ARS)2. In ARS animals randomly forage for food3, but appetitive sensory cues (like an encounter with food) will cause the animal to restrict their search area by reorienting more frequently, thus increasing the likelihood of food encounters. Conversely, when removed from food, animals will decrease their reorientation rate to increase dispersal1,4. Caenorhabditis elegans and Drosophila melanogaster larvae appear to progressively increase their diffusion constant while foraging off of food by decreasing their rates of reorientation4. However in separate studies5,6, individual worms appear to make an abrupt change from a high to low rate of reorientation. This behavior has been described as a switch from a local to global search strategy, that relies on evidence accumulation to trigger the behavioral switch6.

A central challenge with defining these possible state transitions is the stochastic nature of the behavior itself. Reorientations are random and follow Poisson statistics7–13. This is very similar to the temporal behavior of individual molecules in solution14. Individual molecules defined by the same reaction kinetics can stochastically produce long or short time intervals between reaction events, not because one molecule is inherently faster than the other, but because they are stochastically sampling from the same probability distribution. The diversity of times between individual worm reorientations is very similar to this. The exponentially decaying reorientation rate emerges from the average of trajectories that do not necessarily conform to this curve individually. However, even those that produce abrupt switches in reorientation rates could still emerge from a simple exponential decay strategy. Since reorientations occur stochastically, the abrupt changes in reorientation rates could simply be the result of stochastic sampling of an underlying decay phenomenon15. Here, we show that a simple, stochastic model that does not rely on switches in behavioral state is sufficient to reproduce the reorientation kinetics observed in experimental data.

Results

Local-to-global behavioral transitions are inconsistently observed

Earlier version: In López-Cruz et al5., the foraging behaviors of individual worms were tracked for 45 minutes after being removed from food. As observed previously4, the reorientation rate from this study followed an exponential decay (Figure 1a). It was reported that roughly half the worms appeared to make a sudden single switch from local to global search as observed in Calhoun et al6, however the other half appeared to produce no discernable change in search strategy, or exhibited multiple switches (Figure 1 b–d). Whether or not a worm performed a single decision was defined in López-Cruz et al5 by fitting individual reorientation data to two lines using the MATLAB function findchangepts (Figure 1b)16. This function divides each trace into two regions that are defined by minimizing the sum of the residual squared error of two local linear regressions. The location of the transition point is varied until the total residual error attains a minimum (Figure 1b). A large change in slope indicates a sudden change in reorientation rate, and the intersect between the two lines determines the decision-time5 (Figure 1b). A worm was identified as making a sudden local-to-global foraging switch based on visually assessing the magnitude of the slope difference. To evaluate this property more objectively, we assessed the distributions of slope-differences (s1-s2) and transition times to see whether two densities were clearly distinguishable. The resulting distributions for the experimental data were continuous, with no clear boundary between deciders and nondeciders (Figure 1e).

Revised version: In López-Cruz et al.5, the foraging behaviors of individual worms were tracked for 45 minutes after being removed from food. As observed previously4, the reorientation rate from this study followed an exponential decay (Figure 1a). It was reported that roughly half the worms appeared to make a sudden single switch from local to global search as observed in Calhoun et al6, however the other half appeared to produce no discernable change in search strategy, or exhibited multiple switches (Figure 1 b–d). Whether or not a worm performed a single decision was defined in Lopez-Cruz et al5 by fitting individual reorientation data to two lines using the MATLAB function findchangepts (Figure 1b)16. This function divides each trace into two regions that are defined by minimizing the sum of the residual squared error of two local linear regressions. The location of the transition point is varied until the total residual error attains a minimum (Figure 1b). A large change in slope indicates a sudden change in reorientation rate, and the intersect between the two lines determines the decision-time5 (Figure 1b). A worm was identified as making a sudden local-to-global foraging switch based on visually assessing the magnitude of the slope difference, therefore the conclusion that 50% of the worms produced a single transition is fairly subjective. To evaluate this property more objectively, we assessed the distributions of slope-differences (s1-s2) and transition times to see whether two densities were clearly distinguishable. The resulting distributions for the experimental data were continuous, with no clear boundary between deciders and non-deciders (Figure 1e).

Added in the revised version: Foraging kinetics of C. elegans

Added in the revised version: (a) Average experimental population reorientation rate (black line) in a rolling 2-minute window. Blue bins represent probability of observed reorientation rate. (b) Abrupt transitions were identified by performing two linear regressions on observed reorientation curves. Transition times (trans) were defined by the intersection of the regressions (dashed line). The slopes of the two regressions are s1 and s2. (c) An example of an experimental reorientation curve with an abrupt reorientation transition (marked by vertical line). (d) An example of an experimental reorientation curve that lacked an abrupt reorientation transition. (e) Distribution of slope differences and transition times from regressions fit to the experimental data. Individual data points are individual reorientation data for each worm. Insets are individual examples of experimental cumulative reorientation curves. Number of worms (N) = 1631. Dashed line represents the median slope difference. All data curated from López-Cruz et al5.

Why do some worms appear to make a decision, while others do not? In aggregate, the reorientation rate decays (Figure 1a). Despite the population average conforming to a gradual decay, individual trajectories produce a wide diversity of trajectories which sometimes conform to an apparent drop in reorientation rate (Figure 1c), while others do not (Figure 1d). If the worms are executing a decision, this would seem to indicate only a fraction of the worms decide to switch from local to global foraging strategies, while others use a different strategy. An alternative hypothesis is that sudden changes in reorientation are random events, that do not occur due to an underlying change in strategy, but because of the inherently random nature of the behavior itself 7–13.

A stochastic model generates abrupt and gradual changes in search strategy

We tested this hypothesis by modeling individual worms by stochastic sampling of a decaying reorientation rate with the Gillespie algorithm (Figure 2a), a common strategy used to model the kinetics of individual molecules17. With this strategy, the time between chemical events is modeled by randomly sampling from the time-interval distributions defined by the reaction rates. Although the algorithm was originally developed to model discrete molecular events based on known kinetic parameters, it can be used to generate time trajectories for any discrete events when the kinetics are known. A behavioral example of this is the Lotka-Voltera predator-prey competition model where predator and prey populations fluctuate out of phase due to predation. Stochastic fluctuations of predator and prey populations can be modeled using the Gillespie algorithm18–20.

Added in the revised version: Stochastic modeling of foraging kinetics of C. elegans

Added in the revised version: (a) Outline of the Gillespie algorithm. Step 1: two random numbers (r1 and r2) are drawn from a uniform distribution [0,1]. Step 2: The time interval for a new event (τ) is randomly assigned based on the total propensities (a0) and r1. Step 3: The event i (either a decay of M (M←M-1) or a new reorientation (Ω←Ω+1)) at t + τ is determined by the probability (P(i)) of the event i occurring. This probability is determined by the relative propensity (ai/a0) of the event i. (b) (Top) The parameters α, β, and γ are assigned based on fitting a decay curve (red) to the observed average reorientation rate (blue). Dashed lines are + and − standard deviation. (Bottom) Average model population reorientation rate (black line) in a rolling 2-minute window. Red bins represent probability of observed reorientation rate. N = 1631, α = 1.49 min−1, β = 0.1937 min−1, γ = 0.11 min−1. (c) (Top) An example of a modeled abrupt reorientation transition. (Bottom) An example of a modeled reorientation curve that lacked an abrupt reorientation transition. Red data are cumulative reorientations and orange data are M. (d) Distribution of slope differences and transition times from regressions fit to the experimental (blue) and modeled (red) data. Individual data points are individual reorientation data for experimental (blue) or in silico (red) worms. Insets are individual examples of experimental and modeled cumulative reorientation curves. Dashed line represents the median experimental slope difference. (e) Examples of experimental (blue) and modeled (red) cumulative reorientation curves for individual worms, with similar stochastic dynamics. For each example drawn from experiments (blue), the in silico worm with the highest correlation from N=1631 iterations is shown (red). Sudden changes in rate are indicated with arrows. M for each example is shown in orange. (f) Examples of modeled data when the reorientation rate is constant (α = 1.5). Sudden changes in rate are indicated with arrows.

Earlier version: Experimentally, reorientation rate is measured as the number of reorientation events that occurred in an observational window. However, these are discrete stochastic events, so we can describe them in terms of propensity, i.e. the probability of observing a transitional event (in this case, a reorientation) is: dP(Ω+1,t)dt=a1P(Ω,t) (1)

Revised version: Experimentally, reorientation rate is measured as the number of reorientation events that occurred in an observational window. However, these are discrete stochastic events, so we can describe them in terms of propensity, i.e. the probability of observing a transitional event in an infinitesimal time interval dt (in this case, a reorientation) is: dP(Ω+1,t)dt=a1P(Ω,t) (1)

Earlier version: Here, Ω is reorientation number, P(Ω + 1, t) is the probability of observing a reorientation event at time t, and a1 is the propensity for this event to occur. Observationally, the frequency of reorientations observed decays over time, so we can define the propensity as: a1=αe−γt+β (2)

Revised version: Here, Ω is cumulative reorientation number, P(Ω + 1, t) is the probability of observing Ω + 1 cumulative reorientations at time t (i.e. the probability of observing the number of reorientations advance from Ω to Ω +1), and a1 is the propensity for this event to occur, i.e. the likelihood that a particular reaction will occur in the next infinitesimal time interval. In bulk chemical kinetics, this is analogous to the rate of the reaction. Observationally, the frequency of reorientations observed decays over time. This means that the propensity is not constant, so we can define the propensity as decaying in time: a1=αe−γt+β (2)

Where Ω + β is the initial propensity at t=0, and β is the propensity at t = ∞.

Earlier version: We can model this decay as the reorientation propensity coupled to a decaying factor (M): dp(M−1,t)dt=a2P(M,t) (3)

Revised version: The propensities for the Gillespie algorithm are state-dependent; the algorithm is modeling discrete events based on the current state of the system. A time-varying propensity implies it is coupled to a state variable that is changing in time. Since the propensity a1 decays exponentially, it implies it is coupled to a first-order decay process. We can model this decay as the reorientation propensity coupled to a decaying factor (M): dP(M−1,t)dt=a2P(M,t) (3)

Earlier version: Where the propensity of this event (a2) is: a2=−γM (4)

Revised version: Where the propensity of this event (a2) is a first-order decay: a2=−γM (4)

Since M is a first-order decay process, when integrated, the M observed at time t is: M(t)=M(0)e−γt (5)

We can couple the probability of observing a reorientation to this decay by redefining a1 as: a1=αM0M+β (6)

Earlier version: So that now: dp(Ω+1,t)dt=(αe−γt+β)P(Ω,t) (7)

Revised version: In this way, the propensity a1 is explicitly tied to the state M, which is decaying, so that now: dP(Ω+1,t)dt=(αe−γt+β)P(Ω,t) (7)

Earlier version: A critical detail should be noted. While reorientations are modeled as discrete events, the amount of M at time t=0 is chosen to be large (M0 ← 1,000), so that over the timescale of 40 minutes, the decay in M is practically continuous. This ensures that sudden changes in reorientations are not due to sudden changes in M, but due to the inherent stochasticity of reorientations.

Revised version: A critical detail should be noted. While reorientations are modeled as discrete events, the amount of M at time t =0 is chosen to be large (M0 ← 1,000), so that over the timescale of 40 minutes, the decay in M is practically continuous. This ensures that sudden changes in reorientation rate are not due to sudden changes in M, but due to the inherent stochasticity of reorientations.

To model both processes, we can create the master equation: dP(Ω,M,t)dt=a1P(Ω−1,M,t)−a1P(Ω,M,t)+a2P(Ω,M+1,t)−a2P(Ω,M,t) (8)

Earlier version: Since these are both Poisson processes, the probability density function for a state change i occurring at time t is: P(t|ai)=aie−ait (9)

Revised version: Since these are both Poisson processes, the probability density function for a state change i (Ω+1: i =1, M-1: i=2) occurring at time t is: P(t|ai)=aie−ait (9)

Earlier version: The probability that an event will not occur in time interval t is: P(t>τ|ai)=∫τ∞P(t|ai)dt=e−aiτ (10)

Revised version: The probability that a state change will not occur for propensity ai in time interval τ is: P(t>τ|ai)=∫τ∞P(t|ai)dt=e−aiτ (10)

Earlier version: The probability that no events will occur for ALL transitions in this time interval is: P(t>τ|a1)P(t>τ|a2)=∏ie−aiτ=exp[−τ∑iai] (11)

Revised version: The probability that no state changes will occur for ALL propensities in this time interval is: P(t>τ|a1)P(t>τ|a2)=∏ie−aiτ=exp[−τ∑iai] (11)

Earlier version: We can draw a random number (r0 ∈ [0,1]) that represents the probability of no events in time interval τ, so that this time interval can be assigned by rearranging equation 11: τ←−ln(r1)a0 (12)

Revised version: We can draw a random number (r1 ∈ [0,1]) that represents the probability of no events in time interval τ, so that this time interval can be assigned by rearranging equation 11: τ←−ln(r1)a0 (12)

where: a0=∑iai (13)

This is the time interval for any event (Ω+1 or M-1) happening at t + τ. The probability of which event occurs is proportional to its propensity: P(event i)=aia0 (14)

We can draw a second number (r2 ∈ [0,1]) that represents this probability so that which event occurs at time t + τ is determined by the smallest n that satisfies: r2⋅a0≤∑i=1nai (15)

so that: event i←min{n|r2⋅a0≤∑i=1nai} (16)

For example, if the propensity a1 is 10% of a0 at time t, then there is a 10% chance that the resulting reaction will occur at time t. The elegant efficiency of the Gillespie algorithm is two-fold. First, it models all transitions simultaneously, not separately. Second, it provides floating-point time resolution. Rather than drawing a random number, and using a cumulative probability distribution of interval-times to decide whether an event occurs at discrete steps in time, the Gillespie algorithm uses this distribution to draw the interval-time itself. The time resolution of the prior approach is limited by step size, whereas the Gillespie algorithm’s time resolution is limited by the floating-point precision of the random number that is drawn.

In this way, the Gillespie algorithm enables the simultaneous modeling of numerous reactions in parallel with machine-precision time resolution. To ensure the reorientation rate itself is a smooth decay, a large value of M is used to approximate a continuous function. The underlying assumptions are:

Reorientations are stochastic7–13.

All the worms experience a common decay of a signaling factor (M) (Equation 5) that influences the reorientation rate (Equation 7).

Earlier version: In our approach, we fit the exponential curve in Eq. 2 to the average reorientation rate from the experimental data from López-Cruz5 (Figure 2b), and then used these parameters (α, β, γ) to model 1,631 individual worms (the same number of animals in the experimental data). Even though the initial average rates of reorientation are ~1.5 events/minute, there is considerable variance in observed initial rates (Figure 1a). To address this variance, starting values of M were assigned based on randomly drawing a rate from the observed experimental rates at t = 0, and adjusting M so that the starting rate for that in silico worm matched the initial rate of the randomly drawn experimental worm: Mi←M0(r−β)α (17)

Revised version: In our approach, we fit the exponential curve in Eq. 2 to the average reorientation rate from the experimental data from López-Cruz5 (Figure 2b), and then used these parameters (α, β, γ) to model 1,631 individual worms (the same number of animals in the experimental data). Even though the initial average rates of reorientation are ~1.5 events/minute, there is considerable variance in observed initial rates (Figure 1a). To address this variance, starting values of M were assigned based on randomly drawing a rate from the observed experimental rates at t = 0, and adjusting M so that the starting rate for that in silico worm j matched the initial rate of the randomly drawn experimental worm: Mj←M0(r−β)α (17)

where r is an initial rate randomly drawn from the experimental data.

In silico reorientation curves generated produced single-transition and multi/absent transition curves, as observed experimentally (Figure 2c). When plotted along with the experimental data, the in silico data produced a distribution of linear regression parameters comparable to the experimental worms (Figure 2d). The simulation was able to produce individual trajectories that demonstrated switching behavior, despite the lack of a switching mechanism in the model (Equations 2 and 3). Furthermore, our model demonstrated a continuum of switching to non-switching behavior that was observed in experimental results (Figure 2d).

Earlier version: The experimental transition distribution was shifted slightly earlier, however it is important to note that the experimental data were drawn from experiments where 10-15 animals were tracked together. Collisions occurred, but were excluded from analysis. However, in addition to touch, animals also reorient in response to pheromones when they cross the tracks left by other animals, which is more likely early in the recording when the animals are in close proximity 21,22. This in turn may contribute to a higher initial rate early in the experimental observations. Since this dataset does not have information about when and where worms crossed pheromone tracks from other animals, this effect cannot be excluded.

Revised version: The experimental transition distribution was shifted slightly earlier, however it is important to note that the experimental data were drawn from experiments where 10-15 animals were tracked together. Collisions occurred, but were excluded from analysis, and so would not contribute to this observation. However, in addition to collisions, animals also reorient in response to pheromones when they cross the tracks left by other animals, which is more likely early in the recording when the animals are in close proximity 21,22. This in turn may contribute to a higher initial rate early in the experimental observations, which would delay the observed decay in reorientations. Since this dataset does not have information about when and where worms crossed pheromone tracks from other animals, this effect cannot be excluded.

Earlier version: This exception aside, modeling worm foraging behavior with a simple exponential decay of reorientations was sufficient to capture most experimentally observed dynamics, both switch and non-switch-like. Sudden switches between fast and slow reorientation rates were observed in both the experimental and modeled data (Figure 2e), despite the model not relying on a switch strategy, and the decay in M being continuous.

Revised version: This exception aside, modeling worm foraging behavior with a simple exponential decay of reorientations was sufficient to capture most experimentally observed dynamics, both switch and non-switch-like. Sudden switches between fast and slow reorientation rates were observed in both the experimental and modeled data (Figure 2e), despite the model not relying on a switch strategy, and the decay in M being continuous. Sometimes the experimental data produced switch-like behavior (Figure 2e, upper left panel), while other times the reorientation rate appeared to be constant (Figure 2e, lower right panel). The stochastic model produced reorientation data with similar dynamics, despite not relying on a decision paradigm.

Removed in the revised version: Sometimes the experimental data produced switch-like behavior (Figure 2e, upper left panel), while other times the reorientation rate appeared to be constant (Figure 2e, lower right panel). The stochastic model produced reorientation data with similar dynamics, despite not relying on a decision paradigm.

It is important to emphasize that the model was not explicitly designed to match the sudden changes in reorientation rates observed in the experimental data. Kinetic parameters were simply chosen to match the average population behavior. Sudden changes in reorientation rate were not due to sudden changes in the underlying model; stochastic behavior naturally produces sudden bunching of random events. Even if the reorientation rate is set to a constant value, stochastic sampling will still produce sudden changes in rate, even though the rate has no time-dependence (Figure 2f).

Discrete changes in M are inconsistent with experimental data

A large initial value for M0 (M0 = 1000) was initially chosen to ensure a smooth change in reorientation rate, so that sudden changes in reorientation number were not necessarily due to sudden changes in M. To test how sensitive the model was to stochastic changes in M, the initial value of M was tested at M0 = 1000, 100, 10, or 1. On average, all values of M0 were sufficient to reproduce the average reorientation decay observed in experimental data (Figure 3a). The distributions of transition times and slope differences were also comparable for M0 = 1000, 100, or 10, indicating that the model is not particularly sensitive to initial values of M0.

Added in the revised version: Scaling of foraging kinetics of C. elegans

Added in the revised version: (a) Average reorientation rates for experiments (blue), and models where M0 = 1000 (red), M0 = 100 (plum), M0 = 10 (magenta), or M0 = 1 (cyan). Dashed lines indicate standard deviation. (b) Distribution of slope differences and transition times from regressions fit to the experimental data (blue), and models where M0 = 1000 (red), M0 = 100 (plum), M0 = 10 (magenta), or M0 = 1 (cyan). Number of worms for both experiment and models (N) = 1631. Dashed line represents the median slope difference for the experimental data. (c) The Jensen-Shannon divergence between the experimental distribution for slope difference and transition time in (b), and the distributions generated from models where M0 = 1000 (red), M0 = 100 (plum), M0 = 10 (magenta), or M0 = 1 (cyan). Twenty distributions were generated for each M0 model. P-values calculated using Tukey’s range test: <0.001: ***.

However, despite producing a similar average decay in reorientation rates, the distribution of slope differences was distinctly bimodal for M0 = 1 (Figure 3b). This was not surprising, since when M0 = 1, the model is by definition a two-state system. Either the worm has a reorientation rate of α+β (M = 1), or a reorientation rate of β, (M = 0). The probability of switching from M = 1 to M = 0 increases with a timescale of γ (γ = 0.11 min−1). The bimodal slope and transition time distributions for M0 = 1 were distinctly different than those observed for the experimental or M0 = 1000,100,10 models.

Since the slope and transition time distributions were generated from a recursive algorithm rather than a purely analytical model, the model distribution for these metrics was compared to the experimental distribution using the Jensen-Shannon divergence (JSD): JSD(P||Q)=12D(P||M)+12D(Q||M) (18)

Earlier version: where D is the Kullback-Leibler divergence : D(P||Q)=∑xP(x)ln(P(x)Q(x)) (19)

Revised version: where D is the Kullback-Leibler divergence : D(P||Q)=ExP(x)ln(P(x)Q(x)) (19)

and P(x) and Q(x) are the two probability distributions being compared. M(x) is a mixture distribution of P and Q. This metric provides a useful, parameter-free measure for comparing distributions using mutual information. A short distance between two distributions indicates they are more similar than two distributions separated by a longer distance.

To compare how similar the different models (M0 = 1000, 100, 10, 1) compared to the experimental distribution of transition times and slope differences, we ran the model twenty times for each M0, and calculated the JSD between the model and experimental distributions (Figure 3c). The non-binary models (M0 ≠ 1) all produced comparable distances, but the binary models (M0 = 1) were significantly more distant than the rest, indicating this model was more dissimilar than the others when compared to the experimental data. A continuous decline in reorientation rate is more consistent with experimental observations than a discrete change in reorientation rate.

Discussion

The lack of a decision simplifies the dispersal strategy for C. elegans. Rather than relying on the accumulation of evidence to make a discrete decision, the worm relies on a decaying signal in the absence of food that drives the reorientation rate. This strategy increases the diffusion constant of the worm, and ensures a more efficient search strategy to find food4. The stochastic nature of this search is consistent with prior characterizations of worm behavior and neuronal dynamics7–13, and implies that any individual worm would exhibit all the variability of the population if allowed to perform this task multiple times. We consider this to be a null model: simple stochastic models should be sufficient to explain observably stochastic behaviors.

Earlier version: The decay in M can be considered the memory timescale of the last food encounter. With an observed γ of 0.07 min−1, this would mean a memory τ1/2 of ~10 minutes. In principle M should rise on the presence of food to increase the reorientation rate, on a timescale that is not addressed here. The reliance of the reorientation rate on a decaying factor (M) implies that the rate is influenced by a signaling factor which decays in time. It is well established that the reorientation rate is influenced by numerous signaling factors, and is the basis of the biased random walk that drives taxis in shallow gradients7–13.

Revised version: The decay in M can be considered the memory timescale of the last food encounter. With an observed γ of 0.07 min−1, this would mean a memory τ1/2 of ~10 minutes. In principle M should rise on the presence of food to increase the reorientation rate, on a timescale that is not addressed here. In Lopez-Cruz et al5, the starting reorientation rate correlated with food density prior to food-removal, which would be consistent with the amplitude of M correlating with food density/quality. The reliance of the reorientation rate while foraging off of food on a decaying factor (M) implies that the rate is influenced by a signaling factor which decays in time. It is well established that the reorientation rate is influenced by numerous signaling factors, and is the basis of the biased random walk that drives taxis in shallow gradients7–13.

Earlier version: A decay in reorientation rate, rather than a sudden change, is consistent with observations made by López-Cruz et al5. They found that the neurons AIA and ADE redundantly suppress reorientations, and that silencing either one was sufficient to restore the large number of reorientations during early foraging. AIA and ADE synapse onto several neurons that drive reorientations5,23, and their output was inhibited over long timescales (tens of minutes) by presynaptic glutamate binding to MGL-1, a slow G-Protein coupled receptor expressed in AIA and ADE. Their results support a model where sensory neurons suppress the synaptic output of AIA and ADE, which in turn leads to a large number of reorientations early in foraging. As time passes, glutamatergic input from the sensory neurons decrease, which leads to disinhibition of AIA and ADE, and a subsequent suppression of reorientations.

Revised version: A decay in reorientation rate, rather than a sudden change, is consistent with observations made by López-Cruz et al5. They found that the neurons AIA and ADE redundantly suppress reorientations, and that silencing either one was sufficient to restore the large number of reorientations during early foraging. This implies that the timescale of the variable M may represent the timescales of AIA and ADE activities. AIA and ADE synapse onto several neurons that drive reorientations5,23, and their output was inhibited over long timescales (tens of minutes) by presynaptic glutamate binding to MGL-1, a slow G-Protein coupled receptor expressed in AIA and ADE. Their results support a model where sensory neurons suppress the synaptic output of AIA and ADE, which in turn leads to a large number of reorientations early in foraging. As time passes, glutamatergic input from the sensory neurons decrease, which leads to disinhibition of AIA and ADE, and a subsequent suppression of reorientations.

The sensory inputs into AIA and ADE are sequestered into two separate circuits, with AIA receiving chemosensory input and ADE receiving mechanosensory input. Since the suppression of either AIA or ADE is sufficient to increase reorientations, the decay in reorientations is likely due to the synaptic output of both of these neurons decaying in time. This correlates with an observed decrease in sensory neuron activity as well, so the timescale of reorientation decay could be tied to the timescale of sensory neuron activity, which in turn is influencing the timescale of AIA/ADE reorientation suppression. This implies that our factor “M” is likely the sum of several different sensory inputs decaying in time.

The molecular basis of which sensory neuron signaling factors contribute to decreased AIA and ADE activity is made more complicated by the observation that the glutamatergic input provided by the sensory neurons was not essential, and that additional factors besides glutamate contribute to the signaling to AIA and ADE. In addition to this, it is simply not the sensory neuron activity that decays in time, but also the sensitivity of AIA and ADE to sensory neuron input that decays in time. Simply depolarizing sensory neurons after the animals had starved for 30 minutes was insufficient to rescue the reorientation rates observed earlier in the foraging assay. This observation could be due to decreased presynaptic vesicle release, and/or decreased receptor localization on the postsynaptic side.

In summary, there are two neuronal properties that appear to be decaying in time. One is sensory neuron activity, and the other is decreased potentiation of presynaptic input onto AIA and ADE. Our factor “M” is a phenomenological manifestation of these numerous decaying factors. Altered ionotropic glutamate signaling and dopamine release also influence foraging kinetics24, as well as neuropeptides25. The signaling factor M could be the result of one or all of these factors decaying in time. Further work will be needed to reveal how the kinetics of reorientations emerges explicitly from these underlying signaling kinetics.

Removed in the revised version: Foraging kinetics of C. elegans

Removed in the revised version: (a) Average experimental population reorientation rate (black line) in a rolling 2-minute window. Blue bins represent probability of observed reorientation rate. (b) Abrupt transitions were identified by performing two linear regressions on observed reorientation curves. Transition times (trans) were defined by the intersection of the regressions (dashed line). The slopes of the two regressions are s1 and s2. (c) An example of an experimental reorientation curve with an abrupt reorientation transition (marked by vertical line). (d) An example of an experimental reorientation curve that lacked an abrupt reorientation transition. (e) Distribution of slope differences and transition times from regressions fit to the experimental data. Insets are individual examples of experimental cumulative reorientation curves. Number of worms (N) = 1631. Dashed line represents the median slope difference. All data curated from López-Cruz et al5.

Removed in the revised version: Stochastic modeling of foraging kinetics of C. elegans

Removed in the revised version: (a) Outline of the Gillespie algorithm. Step 1: two random numbers (r1 and r2) are drawn from a uniform distribution [0,1]. Step 2: The time interval for a new event (τ) is randomly assigned based on the total propensities (a0) and r1. Step 3: The event i (either a decay of M (M←M-1) or a new reorientation (Ω—Ω+1)) at t + τ is determined by the probability (P(i)) of the event i occurring. This probability is determined by the relative propensity (ai/a0) of the event i. (b) (Top) The parameters α, β, and γ are assigned based on fitting a decay curve (red) to the observed average reorientation rate (blue). Dashed lines are + and − standard deviation. (Bottom) Average model population reorientation rate (black line) in a rolling 2-minute window. Red bins represent probability of observed reorientation rate. N = 1631, α = 1.49 min−1, β = 0.1937 min−1, γ = 0.11 min−1. (c) (Top) An example of a modeled abrupt reorientation transition. (Bottom) An example of a modeled reorientation curve that lacked an abrupt reorientation transition. Red data are cumulative reorientations and orange data are M. (d) Distribution of slope differences and transition times from regressions fit to the experimental (blue) and modeled (red) data. Insets are individual examples of experimental and modeled cumulative reorientation curves. Dashed line represents the median experimental slope difference. (e) Examples of experimental (blue) and modeled (red) cumulative reorientation curves for individual worms, with similar stochastic dynamics. For each example drawn from experiments (blue), the in silico worm with the highest correlation from N=1631 iterations is shown (red). Sudden changes in rate are indicated with arrows. M for each example is shown in orange. (f) Examples of modeled data when the reorientation rate is constant (α = 1.5). Sudden changes in rate are indicated with arrows.

Removed in the revised version: Scaling of foraging kinetics of C. elegans

Removed in the revised version: (a) Average reorientation rates for experiments (blue), and models where M0 = 1000 (red), M0 = 100 (plum), M0 = 10 (magenta), or M0 = 1 (cyan). Dashed lines indicate standard deviation. (b) Distribution of slope differences and transition times from regressions fit to the experimental data (blue), and models where M0 = 1000 (red), M0 = 100 (plum), M0 = 10 (magenta), or M0 = 1 (cyan). Number of worms for both experiment and models (N) = 1631. Dashed line represents the median slope difference for the experimental data. (c) The Jensen-Shannon divergence between the experimental distribution for slope difference and transition time in (b), and the distributions generated from models where M0 = 1000 (red), M0 = 100 (plum), M0 = 10 (magenta), or M0 = 1 (cyan). Twenty distributions were generated for each M0 model. P-values calculated using Tukey’s range test: <0.001: ***.